Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below:

```OCaml
let name = ""
let rollno = ""
```

## Important notes about grading:

1. All code you submit must compile. Programs that do not compile will probably receive an automatic zero. If you are having trouble getting your assignment to compile, please contact the TAs or the instructor or visit the course contact hour. If you run out of time, it is better to comment out the parts that do not compile, than hand in a more complete file that does not compile.
2. All assignments handed in after the deadline will be considered late, and will consume your grace days. 
3. Verify on moodle that you have submitted the correct version, before the deadline. Submitting the incorrect version before the deadline and realizing that you have done so after the deadline will be counted as a late submission.

## Problem 1

Write a function 

```ocaml
cond_drop : 'a list -> ('a -> bool) -> 'a list
```

list that takes in a list and a predicate and drops all elements which satisfy the condition expressed in the predicate. For example, `cond_drop [3;4;5] (fun x -> x mod 2 = 0) = [3;5]`.

<!-- grade_id: cell-7950f1ea0a84b889, points: 0, grade: false, solution: true -->
```OCaml
let rec cond_drop l f =
  (* YOUR CODE HERE *)
  raise (Failure "Not implemented")
```

<!-- grade_id: cell-1d772b02ff64382c, points: 10, grade: true, solution: false -->
```OCaml
(* 10 points *)
assert (cond_drop[3;4;5] (fun x -> x mod 2 = 0) = [3;5])
```

## Problem 2

Write a function 

```ocaml
n_times : ('a -> 'a) -> int -> 'a list -> 'a list
```

such that `n_times (f,n,l)` applies `f` to each element in the list `l` `n` times. For example, `n_times (fun x-> x+1) 50 [0,0] = [50,50]`. If `n<=0` return the input list `l`. Try to write a tail-recursive function. If your function is not tail recursive, there will be a penalty of 2 points. You can use the `List.map` library function.


<!-- grade_id: cell-c28f98bd63f56ba1, points: 0, grade: false, solution: true -->
```OCaml
let rec n_times f n l =
  (* YOUR CODE HERE *)
  raise (Failure "Not implemented")
```

<!-- grade_id: cell-98364653c4b77b2d, points: 10, grade: true, solution: false -->
```OCaml
(* 10 points *)
assert (n_times (fun x-> x+1) 50 [0;0] = [50;50])
```

## Problem 3

Write a function 

```ocaml
range : int -> int -> int -> int list 
```

such that `range num1 num2 step` returns an ordered list of all integers from `num1` to `num2` inclusive which are `step` apart from each other. For example, `range 2 5 1 = [2;3;4;5]`, while `range 2 5 2 = [2;4]`. If `num1 > num2`, then there are no integers in the range, hence you should return the empty list. Raise the exception `IncorrectStep` if `step <= 0`.

<!-- grade_id: cell-1fd147f7eb3b139d, points: 0, grade: false, solution: true -->
```OCaml
exception IncorrectStep

let rec range num1 num2 step =
  (* YOUR CODE HERE *)
  raise (Failure "Not implemented")
```

<!-- grade_id: cell-bba984120a807d42, points: 10, grade: true, solution: false -->
```OCaml
(* 10 points *)
assert (range 2 5 1 = [2;3;4;5]);
assert (range 2 5 2 = [2;4])
```

## Problem 4

Write a function 

```ocaml
zipwith : ('a -> 'b -> 'c) -> 'a list -> 'b list -> 'c list
```

such that `zipwith f l1 l2` generates a list whose `ith` element is obtained by applying `f` to the `ith` element of `l1` and the `ith` element of `l2` . If the lists have different lengths, the extra elements in the longer list should be ignored. For example, `zipwith (+) [1;2;3] [4;5] = [5;7]`.

<!-- grade_id: cell-932054b41dace011, points: 0, grade: false, solution: true -->
```OCaml
let rec zipwith f l1 l2 =
  (* YOUR CODE HERE *)
  raise (Failure "Not implemented")
```

<!-- grade_id: cell-55aefe5fa4f927d0, points: 10, grade: true, solution: false -->
```OCaml
(* 10 points *)
assert (zipwith (+) [1;2;3] [4;5] = [5;7]) 
```

## Problem 5

Write a function 

```ocaml
buckets : ('a -> 'a -> bool) -> 'a list -> 'a list list 
```

that partitions a list into equivalence classes. That is, buckets equiv lst should return a list of lists where each sublist in the result contains equivalent elements, where two elements are considered equivalent if equiv returns true. For example:


```ocaml
buckets (=) [1;2;3;4] = [[1];[2];[3];[4]]
buckets (=) [1;2;3;4;2;3;4;3;4] = [[1];[2;2];[3;3;3];[4;4;4]]
buckets (fun x y -> (=) (x mod 3) (y mod 3)) [1;2;3;4;5;6] = [[1;4];[2;5];[3;6]]
```

The order of the buckets must reflect the order in which the elements appear in the original list. For example, the output of `buckets (=) [1;2;3;4]` should be `[[1];[2];[3];[4]]` and not `[[2];[1];[3];[4]]` or any other permutation.

The order of the elements in each bucket must reflect the order in which the elements appear in the original list. For example, the output of `buckets (fun x y -> (=) (x mod 3) (y mod 3)) [1;2;3;4;5;6]` should be `[[1;4];[2;5];[3;6]]` and not `[[4;1];[5;2];[3;6]]` or any other permutations.

Assume that the comparison function `('a -> 'a -> bool)` is commutative, associative and idempotent.

Just use lists. Do not use sets or hash tables.

List append function `@` may come in handy. `[1;2;3] @ [4;5;6] = [1;2;3;4;5;6]`.

<!-- grade_id: cell-f4d88f256cc9db8b, points: 0, grade: false, solution: true -->
```OCaml
let buckets p l =
  (* YOUR CODE HERE *)
  raise (Failure "Not implemented")
```

<!-- grade_id: cell-fdbd3b5564e151cb, points: 20, grade: true, solution: false -->
```OCaml
(* 20 points *)
assert (buckets (=) [1;2;3;4] = [[1];[2];[3];[4]])
; assert (buckets (=) [1;2;3;4;2;3;4;3;4] = [[1];[2;2];[3;3;3];[4;4;4]])
; assert (buckets (fun x y -> (=) (x mod 3) (y mod 3)) [1;2;3;4;5;6] = [[1;4];[2;5];[3;6]])
```

## Problem 6

A list with consecutive repeated elements can be efficiently encoded by removing the repetitions and instead storing the number of times the element is repeated consecutively. For example, the list `["a";"b";"b";"c";"a";"a";"a";"d";"d";"b";"b"]` can be encoded by the list `[(1,"a");(2,"b");(1,"c");(3,"a");(2,"d");(2,"b")]`. Each tuple counts the number of repetitions of the element. Note that the ordering of elements in the list is preserved in the encoding. Write a function to generate this encoding.

```ocaml
encode : 'a list -> (int * 'a) list
```

Hint: `option` type may be useful here.

<!-- grade_id: cell-8865e0dd05cbd32d, points: 0, grade: false, solution: true -->
```OCaml
let encode (l: 'a list) =
  (* YOUR CODE HERE *)
  raise (Failure "Not implemented")
```

<!-- grade_id: cell-aa3fb845581c6fce, points: 15, grade: true, solution: false -->
```OCaml
(* 15 points *)
assert (encode ["a";"b";"b";"c";"a";"a";"a";"d";"d";"b";"b"] = [(1,"a");(2,"b");(1,"c");(3,"a");(2,"d");(2,"b")]) 
```

## Problem 7

Write a function

```ocaml
flatten : 'a list list -> 'a list
```

that flattens a list. For example, `flatten [[1;2];[3;4]] = [1;2;3;4]`.

<!-- grade_id: cell-1ba30bbe74fd3005, points: 0, grade: false, solution: true -->
```OCaml
let rec flatten l =
  (* YOUR CODE HERE *)
  raise (Failure "Not implemented")
```

<!-- grade_id: cell-8e986c3f5dbf5e89, points: 10, grade: true, solution: false -->
```OCaml
(* 10 points *)
assert (flatten ([[1;2];[3;4]]) = [1;2;3;4])
```

## Problem 8

Write a function

```ocaml

type 'a tree = Leaf | Node of 'a tree * 'a * 'a tree

type traversal = Preorder | Inorder | Postorder

fold_tree : ('a -> 'b -> 'a) -> 'a -> 'b tree -> traversal -> 'a
```

That does a fold of the tree, while visiting the nodes according to the given traversal. For example,

`fold_tree (fun acc x -> acc @ [x]) [] (Node (Node (Leaf,1,Leaf), 2, Node (Leaf,3,Leaf))) Inorder = [1;2;3]`

`fold_tree (fun acc x -> acc @ [x]) [] (Node (Node (Leaf,1,Leaf), 2, Node (Leaf,3,Leaf))) Preorder = [2;1;3]`

`fold_tree (fun acc x -> acc @ [x]) [] (Node (Node (Leaf,1,Leaf), 2, Node (Leaf,3,Leaf))) Postorder = [1;3;2]`

<!-- grade_id: cell-7170388040d2ef4b, points: 0, grade: false, solution: true -->
```OCaml
type 'a tree = Leaf | Node of 'a tree * 'a * 'a tree

type traversal = Preorder | Inorder | Postorder

let rec fold_tree f acc t order =
  (* YOUR CODE HERE *)
  raise (Failure "Not implemented")
```

<!-- grade_id: cell-8bd22519c4af64bb, points: 15, grade: true, solution: false -->
```OCaml
(* 15 points *)
assert (fold_tree (fun acc x -> acc @ [x]) [] (Node (Node (Leaf,1,Leaf), 2, Node (Leaf,3,Leaf))) Inorder = [1;2;3]);
assert (fold_tree (fun acc x -> acc @ [x]) [] (Node (Node (Leaf,1,Leaf), 2, Node (Leaf,3,Leaf))) Preorder = [2;1;3]);
assert (fold_tree (fun acc x -> acc @ [x]) [] (Node (Node (Leaf,1,Leaf), 2, Node (Leaf,3,Leaf))) Postorder = [1;3;2])
```